# D2 : Data Wrangling

In this week's workbook, you'll work through an example in this to gain practice on how to:
* read a DataFrame into Python
* wrangle the data
* clean the data

# Part I: Setup

Data wrangling often requires additional functionality outside what's included in Python by default. For this, we'll import other functions from helpful packages.

**Import the following packages using their common shortened name found in parentheses:**

* `numpy` (`np`)
* `pandas` (`pd`)

<!-- grade_id: cell-b689d21f629d9dda, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
%matplotlib inline
import pandas as pd
import numpy as np
```

<!-- grade_id: cell-2935019a026ca3a1, points: 0.2, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

**Run the following cell code to make things throughout the rest of this workbook a little prettier.** (Note: You don't have to edit code here, but are free to and see what changes to be sure you understand each line.)

```python
# Configure libraries

# Don't display too many rows/cols of DataFrames
pd.options.display.max_rows = 7
pd.options.display.max_columns = 8

# Round decimals when displaying DataFrames
pd.set_option('display.precision', 2)
```

**Read the CSV file at the URL https://raw.githubusercontent.com/fivethirtyeight/data/master/steak-survey/steak-risk-survey.csv into Python in and assign it to the variable `survey`**.

<!-- grade_id: cell-65e31c6985ba9dd2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
survey = pd.read_csv('https://raw.githubusercontent.com/fivethirtyeight/data/master/steak-survey/steak-risk-survey.csv')

# take a look at the data
survey
```

Output:
```
RespondentID  \
0             NaN   
1        3.24e+09   
2        3.23e+09   
..            ...   
548      3.23e+09   
549      3.23e+09   
550      3.23e+09   

    Consider the following hypothetical situations: <br>In Lottery A, you have a 50% chance of success, with a payout of $100. <br>In Lottery B, you have a 90% chance of success, with a payout of $20. <br><br>Assuming you have $10 to bet, would you play Lottery A or Lottery B?  \
0                                             Response                                                                                                                                                                                                                                  
1                                            Lottery B                                                                                                                                                                                                                            
```

<!-- grade_id: cell-0fdae0bebda7c436, points: 0.2, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
### BEGIN TESTS
assert survey.shape == (551, 15)
### END TESTS
```

<!-- grade_id: cell-caef12a25ae8fc28, points: 0.2, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

These data contain survey responses from Americans who responded to a SurveyMonkey Audience poll. These data were used in the [FiveThirtyEight](https://fivethirtyeight.com) article: *[How Americans Like Their Steak](https://fivethirtyeight.com/features/how-americans-like-their-steak/)*

# Part II: Wrangling

**Write a line of code to look at the first few rows of the DataFrame**

<!-- grade_id: cell-c7cd7d4cfbf98114, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
survey.head(5)
```

Output:
```
RespondentID  \
0           NaN   
1      3.24e+09   
2      3.23e+09   
3      3.23e+09   
4      3.23e+09   

  Consider the following hypothetical situations: <br>In Lottery A, you have a 50% chance of success, with a payout of $100. <br>In Lottery B, you have a 90% chance of success, with a payout of $20. <br><br>Assuming you have $10 to bet, would you play Lottery A or Lottery B?  \
0                                           Response                                                                                                                                                                                                                                  
1                                          Lottery B                                                                                                                                                                                                                                  
2                                          Lottery 
```

What do you notice about the first row of the dataframe? Notice that it's not actually an observation from a respondent? **Remove this row from the dataset. Assign this back to the variable `survey`. Print the first few rows again to make sure you've accomplished this.**

<!-- grade_id: cell-0624532e0d0d4ce1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
survey = survey.drop(labels=[0], axis='rows')
survey.head()
```

Output:
```
RespondentID  \
1      3.24e+09   
2      3.23e+09   
3      3.23e+09   
4      3.23e+09   
5      3.23e+09   

  Consider the following hypothetical situations: <br>In Lottery A, you have a 50% chance of success, with a payout of $100. <br>In Lottery B, you have a 90% chance of success, with a payout of $20. <br><br>Assuming you have $10 to bet, would you play Lottery A or Lottery B?  \
1                                          Lottery B                                                                                                                                                                                                                                  
2                                          Lottery A                                                                                                                                                                                                                                  
3                                          Lottery 
```

<!-- grade_id: cell-3969e31aae66b836, points: 0.2, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
### BEGIN TESTS
assert survey.shape == (550, 15)
### END TESTS
```

<!-- grade_id: cell-d3cf9af60d6a942b, points: 0.2, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

Notice that there are a lot of different questions that were asked of respondents (columns) and 550 people who responded to the survey (rows). 

**Print a list of all the column names in this DataFrame.**

We'll only end up working with a subset of these.

<!-- grade_id: cell-c91f7a1c6729ad43, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
print(survey.columns)
```

Output:
```
Index(['RespondentID',
       'Consider the following hypothetical situations: <br>In Lottery A, you have a 50% chance of success, with a payout of $100. <br>In Lottery B, you have a 90% chance of success, with a payout of $20. <br><br>Assuming you have $10 to bet, would you play Lottery A or Lottery B?',
       'Do you ever smoke cigarettes?', 'Do you ever drink alcohol?',
       'Do you ever gamble?', 'Have you ever been skydiving?',
       'Do you ever drive above the speed limit?',
       'Have you ever cheated on your significant other?', 'Do you eat steak?',
       'How do you like your steak prepared?', 'Gender', 'Age',
       'Household Income', 'Education', 'Location (Census Region)'],
      dtype='object')
```

Now we have a sense of what information is included in the dataset. In the coming weeks, we'll answer the following questions:
1. Who cheats more on their significant other - males or females?
2. Are cigarette smokers less likely to skydive?
3. Do people in New England gamble more than other parts of the country?

To answer these we'll only need data from *some* of the columns in the dataset.

Let's drop the columns we don't need. **Drop the first two columns from the dataset. This should still be assigned to the variable `survey`.**

<!-- grade_id: cell-8a1106312081bab5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
survey = survey.drop(survey.columns[[0,1]], axis='columns')
survey
```

Output:
```
Do you ever smoke cigarettes? Do you ever drink alcohol?  \
1                             NaN                        NaN   
2                              No                        Yes   
3                              No                        Yes   
..                            ...                        ...   
548                           Yes                        Yes   
549                            No                        Yes   
550                           Yes                        Yes   

    Do you ever gamble? Have you ever been skydiving?  ...    Age  \
1                   NaN                           NaN  ...    NaN   
2                    No                            No  ...   > 60   
3                   Yes                            No  ...   > 60   
..                  ...                           ...  ...    ...   
548                 Yes                            No  ...  30-44   
549                 Yes                            No  ...  45-60   
550     
```

<!-- grade_id: cell-89822f596640cd4e, points: 0.2, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
### BEGIN TESTS
assert survey.shape == (550,13)
### END TESTS
```

<!-- grade_id: cell-7951ab60702eea97, points: 0.2, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

Now that we've got the columns we want, let's clean up those column names. **Rename the columns in `survey` so the appropriate columns have the following names:**

* smoking
* alcohol
* gambling
* skydiving
* speeding
* cheated
* steak
* steak_preference
* gender
* age
* income 
* education
* region

Note that a list has been provided here, with these names, so you don't have to type them all:

```python
['smoking','alcohol','gambling',
 'skydiving','speeding', 'cheated',
 'steak', 'steak_preference','gender',
 'age', 'income', 'education', 'region']
```

Output:
```
['smoking',
 'alcohol',
 'gambling',
 'skydiving',
 'speeding',
 'cheated',
 'steak',
 'steak_preference',
 'gender',
 'age',
 'income',
 'education',
 'region']
```

```python
['Do you ever smoke cigarettes?', 'Do you ever drink alcohol?',
       'Do you ever gamble?', 'Have you ever been skydiving?',
       'Do you ever drive above the speed limit?',
       'Have you ever cheated on your significant other?', 'Do you eat steak?',
       'How do you like your steak prepared?', 'Gender', 'Age',
       'Household Income', 'Education', 'Location (Census Region)']
```

Output:
```
['Do you ever smoke cigarettes?',
 'Do you ever drink alcohol?',
 'Do you ever gamble?',
 'Have you ever been skydiving?',
 'Do you ever drive above the speed limit?',
 'Have you ever cheated on your significant other?',
 'Do you eat steak?',
 'How do you like your steak prepared?',
 'Gender',
 'Age',
 'Household Income',
 'Education',
 'Location (Census Region)']
```

<!-- grade_id: cell-ff9ae6330821a062, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
survey = survey.rename(columns={'Do you ever smoke cigarettes?':'smoking', 'Do you ever drink alcohol?':'alcohol', 'Do you ever gamble?':'gambling', 'Have you ever been skydiving?':'skydiving',
       'Do you ever drive above the speed limit?':'speeding',
       'Have you ever cheated on your significant other?':'cheated', 'Do you eat steak?':'steak',
       'How do you like your steak prepared?':'steak_preference', 'Gender':'gender', 'Age':'age',
       'Household Income':'income', 'Education':'education', 'Location (Census Region)':'region'})
print(survey.columns)
```

Output:
```
Index(['smoking', 'alcohol', 'gambling', 'skydiving', 'speeding', 'cheated',
       'steak', 'steak_preference', 'gender', 'age', 'income', 'education',
       'region'],
      dtype='object')
```

<!-- grade_id: cell-6f37071b8c831964, points: 0.1, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
### BEGIN TESTS
assert survey.shape == (550,13)
### END TESTS
```

<!-- grade_id: cell-ac6bcfa7d1ac939d, points: 0.1, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
```

We're in pretty good shape now. **Print the first few rows of the `survey` DataFrame to see what the data look like at this point.**

```python
survey.head()
```

Output:
```
smoking alcohol gambling skydiving  ...   age             income  \
1     NaN     NaN      NaN       NaN  ...   NaN                NaN   
2      No     Yes       No        No  ...  > 60  $50,000 - $99,999   
3      No     Yes      Yes        No  ...  > 60          $150,000+   
4     Yes     Yes      Yes        No  ...  > 60  $50,000 - $99,999   
5      No     Yes       No        No  ...  > 60  $50,000 - $99,999   

                          education              region  
1                               NaN                 NaN  
2  Some college or Associate degree  East North Central  
3                   Graduate degree      South Atlantic  
4                   Bachelor degree         New England  
5                   Graduate degree     Middle Atlantic  

[5 rows x 13 columns]
```

# Part III: Cleaning

**Now that we've go the data we need, let's get a sense for how much missing data there is in this dataset by determining how many null-containing rows there are in `survey`. Assign this value to the variable `null_rows`.**

<!-- grade_id: cell-5bc6a10438afe331, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
null_rows=survey.isnull().any(axis='columns').sum()
print(null_rows)
```

Output:
```
217
```

<!-- grade_id: cell-e43974ae43ee28d8, points: 0.2, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
### BEGIN TESTS
assert null_rows == 217
### END TESTS
```

Good to know that lots of people didn't answer every question. We'll keep that in mind as we work with this dataset.

Simply dropping missing observations is typically not good practice; however, in this case we'll drop observations that have missing data across the entire row, as these are individuals who didn't participate in the survey at all. **Remove rows where ALL the columns have missing data for that participant.**

Run the following line to print out the docstring of the `dropna` function which includes description of what this method does, parameters and default values, and some usage examples.

```python
survey.dropna?
```

Output:
```
[0;31mSignature:[0m
[0msurvey[0m[0;34m.[0m[0mdropna[0m[0;34m([0m[0;34m[0m
[0;34m[0m    [0;34m*[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0maxis[0m[0;34m:[0m [0;34m'Axis'[0m [0;34m=[0m [0;36m0[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mhow[0m[0;34m:[0m [0;34m'AnyAll | lib.NoDefault'[0m [0;34m=[0m [0;34m<[0m[0mno_default[0m[0;34m>[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mthresh[0m[0;34m:[0m [0;34m'int | lib.NoDefault'[0m [0;34m=[0m [0;34m<[0m[0mno_default[0m[0;34m>[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0msubset[0m[0;34m:[0m [0;34m'IndexLabel | None'[0m [0;34m=[0m [0;32mNone[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0minplace[0m[0;34m:[0m [0;34m'bool'[0m [0;34m=[0m [0;32mFalse[0m[0;34m,[0m[0;34m[0m
[0;34m[0m    [0mignore_index[0m[0;34m:[0m [0;34m'bool'[0m [0;34m=[0m [0;32mFalse[0m[0;34m,[0m[0;34m[0m
[0;34m[0m[0;34m)[0m [0;34m->[0m [0;34m'DataFrame | None'[0m[0;34m[0m[0;34m[0m
```

<!-- grade_id: cell-b8417c0e50579e72, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
survey = survey.dropna(how='all')
```

<!-- grade_id: cell-27d421b2b5aaac0f, points: 0.2, grade: true, solution: false -->
```python
# Do not edit this cell. This cell is for testing purposes
### BEGIN TESTS
assert survey.shape == (541, 13)
### END TESTS
```

**Print the first few rows to remind yourself what the data look like at this point.**

<!-- grade_id: cell-e1148cbe6b8da30b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
print(survey.head())
```

Output:
```
smoking alcohol gambling skydiving  ...    age             income  \
2      No     Yes       No        No  ...   > 60  $50,000 - $99,999   
3      No     Yes      Yes        No  ...   > 60          $150,000+   
4     Yes     Yes      Yes        No  ...   > 60  $50,000 - $99,999   
5      No     Yes       No        No  ...   > 60  $50,000 - $99,999   
6      No      No       No        No  ...  18-29       $0 - $24,999   

                          education              region  
2  Some college or Associate degree  East North Central  
3                   Graduate degree      South Atlantic  
4                   Bachelor degree         New England  
5                   Graduate degree     Middle Atlantic  
6  Some college or Associate degree  West South Central  

[5 rows x 13 columns]
```

Note that the first row no longer has all missing data here. We've got a dataset we can work with now!

**Great work on this workbook! We'll continue to work with this dataset in section to answer our questions of interest. You can work on other parts of the class or help a classmate work through this - we always understand things best once we've had to explain them to someone else.**

```python

```

```python

```

```python

```

```python

```

```python

```